# Estructuras de datos

Cuando tengamos que trabajar con muchos datos simultáneamente, será conveniente guardarlos en estructuras de datos. <br>

En este notebook cubriremos:
1. Tuplas (Tuples)
2. Diccionarios (Dictionarys)
3. Arreglos (Arrays)

<br>
A grandes rasgos, las tuplas y los arreglos son ordenados (por lo cual podemos usar índices con ellos), y los diccionarios y areglos son mutables.
¡Explicaremos más más abajo!

## Tuplas

Podemos crear una tupla encerrando una colección de elementos entre paréntesis `( )`.

Sintaxis: <br>
```julia
(item1, item2, ...)
```


```julia
misanimalesfavoritos = ("pingüinos", "gatos", "hipopótamos")
```

Output:
```
Tuple{String, String, String}
```

```julia
typeof(misanimalesfavoritos)
```

Output:
```
Tuple{String, String, String}
```

Podemos usar un índice para leer elementos de la tupla.

```julia
misanimalesfavoritos[1]
```

Output:
```
"pingüinos"
```

Las tuplas son inmutables, es decir, una vez creadas no podemos cambiar sus elementos:

```julia
misanimalesfavoritos[1] = "marmota"
```

Output:
```
@ ./boot.jl:373 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1196
```

**ACLARACIÓN IMPORTANTE:** Note que utilizamos el índice `1` para referirnos al primer elemento. En otros lenguajes de programación (por ejemplo Python), una tupla o arreglo de `N` elementos se indexa con los números del `0` al `N-1`. En Julia los índices van de `1` a `N`, independientemente de las preferencias personales del escritor de esta aclaración.

## TuplasNombradas (NamedTuples): ¡Desde Julia 1.6 en adelante!

Las tuplas nombradas son iguales a las tuplas, solo que podemos agregarle un nombre a los elementos.
```julia
(nombre1 = item1, nombre2 = item2, ...)
```
Note que en una tupla nombrada todos los elementos deben tener nombre.

```julia
misanimalesfavoritos = (aves = "pingüino", felinos = "gato", anfibios = "hipopótamos")
```

Output:
```
(aves = "pingüino", felinos = "gato", anfibios = "hipopótamos")
```

Podemos acceder a sus elementos usando índices

```julia
misanimalesfavoritos[1]
```

Output:
```
"pingüino"
```

Pero también podemos usar sus nombres

```julia
misanimalesfavoritos.anfibios
```

Output:
```
"hipopótamos"
```

## Diccionarios

Si tenemos varios sets de datos relacionados entre sí, podemos elegir guardarlos en un diccionario. Podemos crear diccionarios utilizando la función `Dict()`, la cual nos permite inicializar diccionarios vacíos o con pares `key` (clave) `value` (valor).

Sintaxis:
```julia
Dict(key1 => value1, key2 => value2, ...)
```
Podríamos por ejemplo crear un diccionario para guardar una lista de contáctos

```julia
miscontactos = Dict("Angélica" => "867-5309", "Eliza" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Angélica" => "867-5309"
  "Eliza"    => "555-2368"
```

En este ejemplo, cada nombre y número es un par "key" "valor". Podemos leer el número de teléfono de Angélica (valor) usando la "key" asociada:

```julia
miscontactos["Angélica"]
```

Output:
```
"867-5309"
```

Podemos agregar más elementos al diccionario de manera sencilla

```julia
miscontactos["Peggy"] = "555-0234"
```

Output:
```
"555-0234"
```

O modificar entradas existentes

```julia
miscontactos["Angélica"] = "999-1245"
```

Output:
```
"999-1245"
```

Veamos cómo se ve el diccionario ahora

```julia
miscontactos
```

Output:
```
Dict{String, String} with 3 entries:
  "Peggy"    => "555-0234"
  "Angélica" => "999-1245"
  "Eliza"    => "555-2368"
```

Podemos quitar un elemento del diccionario usando `pop!`

```julia
pop!(miscontactos, "Peggy")
```

Output:
```
"555-0234"
```

```julia
miscontactos
```

Output:
```
Dict{String, String} with 2 entries:
  "Angélica" => "999-1245"
  "Eliza"    => "555-2368"
```

A diferencia de los arrays y las tuplas, los diccionarios no están ordenados. Fíjese que las veces que verificamos los contenidos del diccionario escribiendo `miscontactos`, el orden en que aparecieron las entradas fue aleatorio. Por esta razón no podemos usar índices con los diccionarios.

```julia
miscontactos[1]
```

Output:
```
@ ./boot.jl:373 [inlined]
 [4] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1196
```

En este ejemplo, `julia` cree que estás buscando el valor asociado con el key `1`, que no existe.

## Arreglos

Los arreglos son muy similares a las tuplas, pero sus elementos sí se pueden modificar.

Sintaxis: <br>
```julia
[item1, item2, ...]
```

Por ejemplo, puedo crear un arreglo de mis gemas favoritas

```julia
gemas = ["Lapislázuli", "Perla", "Cuarzo Rosa", "Ametista", "Rubí", "Zafiro", "Granate"]
```

Output:
```
7-element Vector{String}:
 "Lapislázuli"
 "Perla"
 "Cuarzo Rosa"
 "Ametista"
 "Rubí"
 "Zafiro"
 "Granate"
```

```julia
typeof(gemas)
```

Output:
```
Vector{String} (alias for Array{String, 1})
```

El `1` de `Array{String,1}` significa que es un Array de una dimensión.  Un `Array{String,2}` sería de dos dimensiones (una matriz), etc. `String` is el tipo de cada elemento.

También podemos guardar números

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

¡O combinar!

```julia
mezcla = [1, 1, 2, 3, "Zafiro", "Rubí"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Zafiro"
  "Rubí"
```

Una vez que tenemos un arreglo, podemos acceder a sus elementos usando índice. Por ejemplo, si queremos ver cual es la tercera gema en la lista `gemas` escribimos

```julia
gemas[3]
```

Output:
```
"Cuarzo Rosa"
```

Podemos usar el índice para cambiar un elemento

```julia
gemas[3] = "Steven"
```

Output:
```
"Steven"
```

```julia
gemas
```

Output:
```
7-element Vector{String}:
 "Lapislázuli"
 "Perla"
 "Steven"
 "Ametista"
 "Rubí"
 "Zafiro"
 "Granate"
```

También podemos editar arrays usando las funciones `push!` y `pop!`. `push!` agrega un elemento al final del array y `pop!` elimina el último elemento de un array.

Podemos agregar un número a la secuencia de fibonacci

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
 21
```

Y luego eliminarlo

```julia
pop!(fibonacci)
```

Output:
```
21
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

Hasta ahora hemos utilizado arreglos de una dimensión, pero estos pueden tener dimensión arbitraria.
<br><br>
Por ejemplo, los siguientes son arrays de arrays

```julia
favoritos = [["Brócoli", "Berenjena", "Espinaca"],["Pingüinos", "Gatos", "Hipopótamos"]]
favoritos[1][2]  #forma de acceder a los elementos de un arreglo de arreglos
```

Output:
```
"Berenjena"
```

```julia
numeros = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Los siguientes dos son ejemplos de arrays bidimensionales y tridimensionales de números aleatorios entre 0 y 1.

```julia
A = rand(4, 3) #matriz 4x3
```

Output:
```
4×3 Matrix{Float64}:
 0.157214   0.894853  0.375661
 0.178866   0.975147  0.754258
 0.0932494  0.924416  0.182343
 0.453078   0.358664  0.867225
```

```julia
A[1,2]     #forma de acceder a los elementos de una matriz A[1][2] dará error
           #note que el primer número indica fila y el segundo indica columna
```

Output:
```
0.8948534152129032
```

```julia
rand(4, 3, 2) #matrix 4x3x2
```

Output:
```
4×3×2 Array{Float64, 3}:
[:, :, 1] =
 0.629816  0.549706  0.614519
 0.904671  0.890294  0.561678
 0.755543  0.926331  0.590984
 0.100227  0.702597  0.660997

[:, :, 2] =
 0.404188  0.83967    0.220091
 0.598368  0.161274   0.824898
 0.413578  0.0236151  0.979353
 0.264105  0.914657   0.974979
```

**¡Cuidado a la hora de copiar arreglos!**
Supongamos que quiero hacer una copia del arreglo `fibonacci`.
Probemos hacerlo escribiendo `otrosnumeros = fibonacci`.

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
otrosnumeros = fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

Por ahora todo parece andar bien. Tenemos aparentemente dos arreglos idénticos. Sin embargo, vea qué sucede con `fibonacci` si modificamos `otrosnumeros`:

```julia
otrosnumeros[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
 404
   1
   2
   3
   5
   8
  13
```

¡Editar `otrosnumeros` también cambió `fibonacci`!.

En el ejemplo anterior no creamos una copia de `fibonacci`. Solo creamos otra forma de acceder al arreglo `fibonacci`. Es decir, para Julia, `otrosnumeros` y `fibonacci` son dos nombres distintos para un mismo objeto.

Si queremos hacer una copia independiente de `fibonacci` podemos hacerlo con la función `copy`

```julia
# Primero, restauramos fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
otrosnumeros = copy(fibonacci)
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
otrosnumeros[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

En este ejemplo, `fibonacci` no cambió cuando modificamos `otrosnumeros`. Por lo tanto ambos arreglos son distintos.

### Algunos trucos

```julia
miarray = ["a", "b", "c", "d", "e", "f", "g", "h", "i"]
```

Output:
```
9-element Vector{String}:
 "a"
 "b"
 "c"
 "d"
 "e"
 "f"
 "g"
 "h"
 "i"
```

```julia
#Cómo acceder al último elemento
miarray[end]
```

Output:
```
"i"
```

```julia
#Cómo acceder al penúltimo elemento
miarray[end-1]
```

Output:
```
"h"
```

```julia
#Cómo leer un array desde el índice `inicio` hasta `fin` saltando de a `salto` elementos
salto = 2
inicio = 1
fin = 6

miarray[inicio:salto:fin]
```

Output:
```
3-element Vector{String}:
 "a"
 "c"
 "e"
```

```julia

```

### Ejercicio

#### 3.1 
Cre un arreglo, `a_ray`, con el siguiente código:

```julia
a_ray = [1, 2, 3]
```

Agregue el número 4 al array y luego elimínelo

```julia
#Cree el array
#Agregue el número 4
a_ray = [1,2,3]
push!(a_ray, 4)
```

Output:
```
4-element Vector{Int64}:
 1
 2
 3
 4
```

```julia
@assert a_ray == [1, 2, 3, 4]
```

```julia
#Elimine el último número

pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Trate de agregar el número "Emergencias" a `miscontactos` con el valor `string(911)` usando el siguiente código
```julia
miscontactos["Emergencias"] = 911
```

¿Por qué esto no funciona?

```julia
miscontactos
```

Output:
```
Dict{String, String} with 3 entries:
  "Emergencias" => "911"
  "Angélica"    => "999-1245"
  "Eliza"       => "555-2368"
```

```julia
miscontactos["Emergencias"] = "911"
```

Output:
```
"911"
```

```julia
miscontactos
```

Output:
```
Dict{String, String} with 3 entries:
  "Emergencias" => "911"
  "Angélica"    => "999-1245"
  "Eliza"       => "555-2368"
```

#### 3.3 
Cree un nuevo diccionario llamado `miscontactos_flexible` que tenga el número de Angélica guardado como un número entero y el de Eliza como un string, usando el código
```julia
miscontactos_flexible = Dict("Angélica" => 867-5309, "Eliza" => "555-2368")
```

```julia
miscontactos_flexible = Dict("Angélica" => 867-5309, "Eliza" => "555-2368")
```

Output:
```
Dict{String, Any} with 2 entries:
  "Angélica" => -4442
  "Eliza"    => "555-2368"
```

```julia
@assert miscontactos_flexible == Dict("Angélica" => 867-5309, "Eliza" => "555-2368")
```

#### 3.4 
Ahora pruebe nuevamente agregar "Emergencias" con el valor `911` (como entero) a `miscontactos_flexible`.

```julia
miscontactos_flexible["Emergencias"] = Int64(911)
```

Output:
```
911
```

```julia
@assert haskey(miscontactos_flexible, "Emergencias")
#haskey(diccionario, key) retorna verdadero si el diccionario tiene una entrada "key" y falso si no la tiene.
```

```julia
@assert miscontactos_flexible["Emergencias"] == Int64(911)
```

#### 3.5 
¿Por qué podemos agregar un entero a `miscontactos_flexible` pero no a `miscontactos`?, ¿Cómo podríamos haber inicializado `miscontactos` para que acepte enteros como valores? (pista: intente usar [la documentación de Julia sobre diccionarios](https://docs.julialang.org/en/v1/base/collections/#Dictionaries)).

```julia

```