# Despacho Múltiple (Multiple dispatch)

En este notebook exploraremos **multiple dispatch**, que es una característica de Julia.

Multiple dispatch ayuda a hacer código genérico y rápido

#### Empezando con lo conocido

Ya hemos visto que podemos declarar funciones en Julia sin dar información sobre los tipos de los argumentos que recibe la función. Por ejemplo, podemos definir

```julia
f(x) = x.^2
```

y Julia determinará solo qué clase de argumentos tienen sentido y que clase no

```julia
f(10)
```

```julia
f([1, 2, 3])
```

#### Especificar los tipos de variable de los argumentos

Sin embargo, también tenemos la *opción* de definir el tipo de variable que recibe una función de Julia.
Por ejemplo, escribamos una función `foo` que solo tome `String`s como argumentos.

```julia
foo(x::String, y::String) = println("¡Tanto x como y son strings!")
```

Vemos que para restringir el tipo de `x` e `y` a `String`, simplemente escribimos inmediatamente después del argumento dos doble puntos seguidos de la palabra `String`.

En los ejemplos siguientes veremos que `foo` solo funciona si tanto `x` como `y` son `String`s.

```julia
foo("hola", "mundo")
```

```julia
foo(3, 4)
```

Para que `foo` funcione con argumentos de tipo `Int` (enteros), declaremos `foo` agregando `::Int` a los argumentos.

```julia
foo(x::Int, y::Int) = println("¡Tanto x como y son enteros!")
```

```julia
foo(3, 4)
```

¡Ahora `foo` funciona con enteros! Pero no solo eso, ¡`foo` también sigue funcionando cuando `x` e `y` son strings!

```julia
foo("hola", "aloha")
```

Esto comienza a acercarse a la idea central de **multiple dispatch**. Cuando declaramos

```julia
foo(x::Int, y::Int) = println("¡Tanto x como y son enteros!")
```
no reescribimos o reemplazamos
```julia
foo(x::String, y::String)
```
Sino que simplemente agregamos un ***método (method)*** a la ***función genérica (generic function)*** llamada `foo`.

Una ***función genérica*** es un concepto abstracto asociado a una operación particular.

Por ejemplo, la función genérica `+` representa el concepto de adición.

Un ***método*** es una implementación específica de una función genérica para *tipos de argumentos particulares*.

Por ejemplo, `+` tiene métodos que aceptan números de punto flotante, enteros, matrices, etc.

Podemos usar la función `methods` para ver cuantos métodos tiene `foo`.

```julia
methods(foo)
```

Nota: ¿cuántos métodos cree que tiene `+`?

```julia
methods(+)
```

Entonces, podemos llamar a `foo` con enteros o strings. Cuando uno llama a `foo` con un conjunto particular de argumentos, Julia inferirá los tipos de argumentos y hará un ***dispatch*** del método apropiado. *Eso* es multiple dispatch.

Multiple dispatch permite tener código genérico y flexible, y a su vez rápido, ya que Julia es capaz de llamar a métodos eficientes para los tipos de argumentos que le demos a las funciones.

Para ver qué métodos usará Julia para ciertos argumentos, podemos usar el macro @which:

```julia
@which foo(3, 4)
```

¡Veamos que pasa si usamos `@which` con el operador de suma!

```julia
@which 3.0 + 3.0
```

Y podemos seguir agregando métodos a  `foo`. Agreguemos uno que tome como argumento un ***tipo abstracto (abstract type)*** `Number`, que incluye subtipos tales como `Int`, `Float64`, y otros objetos que usualmente pensaríamos cono números:

```julia
foo(x::Number, y::Number) = println("¡Mis inputos x e y son números!")
```

Este método de `foo` funcionará, por ejemplo, en números de punto flotante

```julia
foo(3.0, 4.0)
```

También podemos agregar un método de respaldo que tome argumentos de cualquier tipo

```julia
foo(x, y) = println("¡Ahora funciono con cualquier argumento para x o y!")
```

```julia
v = rand(3)
w = "hipopótamo"

foo(v,w)
```

### Ejercicios

#### 9.1

Extienda la función `foo`, agregando un método que acepte un único argumento, y que este argumento sea de tipo `Bool`. Haga que este método retorne "foo con un booleano!"

```julia
foo(x::Bool) = "foo con un booleano!"
```

#### 9.2

Chequee que el método 
```julia
foo(true)
```
sea el que escribió usted

```julia
typeof(foo(true))

```

```julia
@which foo(true)
```

<!-- grade_id: cell-14072e60ae07c1a2, points: 1, grade: true, solution: false -->
```julia
@assert (foo(true) == "foo con un booleano!")
```